Eye-specific differences in active zone addition during synaptic competition in the developing visual system

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Summary

The authors previously published a study of RGC boutons in the dLGN in developing wild-type mice and developing mutant mice with disrupted spontaneous activity. In the current manuscript, they have broken down their analysis of RGC boutons according to the number of Homer/Bassoon puncta associated with each vGlut3 cluster.

The authors find that, in the first post-natal week, RGC boutons with multiple active zones (mAZs) are about a third as common as boutons with a single active zone (sAZ). The size of the vGluT2 cluster associated with each bouton was proportional to the number of active zones present in each bouton. Within the author's ability to estimate these values (n=3 per group, 95% of results expected to be within ~2.5 standard deviations), these results are consistent across groups: 1) dominant eye vs. nondominant eye, 2) wild-type mice vs. mice with activity blocked, and at 3) ages P2, P4, and P8. The authors also found that mAZs and sAZs also have roughly the same number (about 1.5) of sAZs clustered around them (within 1.5 um).

However, the authors do not interpret this consistency between groups as evidence that active zone clustering is not a specific marker or driver of activity dependent synaptic segregation. Rather, the authors perform a large number of tests for statistical significance and cite the presence or absence of statistical significance as evidence that "Eye-specific active zone clustering underlies synaptic competition in the developing visual system (title)". I don't believe this conclusion is supported by the evidence.

We have revised the title to be descriptive: "Eye-specific differences in active zone addition during synaptic competition in the developing visual system." While our correlative approach does not establish direct causality, our findings provide important structural evidence that complements existing functional studies of activity-dependent synaptic refinement. We have carefully revised the text throughout to avoid causal language, focusing instead on the developmental patterns we observe.

Strengths

The source dataset is high resolution data showing the colocalization of multiple synaptic proteins across development. Added to this data is labeling that distinguishes axons from the right eye from axons from the left eye. The first order analysis of this data showing changes in synapse density and in the occurrence of multi-active zone synapses is useful information about the development of an important model for activity dependent synaptic remodeling.

Weaknesses

In my previous review I argued that it was not possible to determine, from their analysis, whether the differences they were reporting between groups was important to the biology of the system. The authors have made some changes to their statistics (paired t-tests) and use some less derived measures of clustering. However, they still fail to present a meaningfully quantitative argument that the observed group differences are important. The authors base most of their claims on small differences between groups. There are two big problems with this practice. First, the differences between groups appear too small to be biologically important. Second, the differences between groups that are used as evidence for how the biology works are generally smaller than the precision of the author's sampling. That is, the differences are as likely to be false positives as true positives.

(1) Effect size. The title claims: "Eye-specific active zone clustering underlies synaptic competition in the developing visual system". Such a claim might be supported if the authors found that mAZs are only found in dominant-eye RGCs and that eye-specific segregation doesn't begin until some threshold of mAZ frequency is reached. Instead, the behavior of mAZs is roughly the same across all conditions. For example, the clear trend in Figure 4C and D is that measures of clustering between mAZ and sAZ are as similar as could reasonably be expected by the experimental design. However, some of the comparisons of very similar values produced p-values < 0.05. The authors use this fact to argue that the negligible differences between mAZ and sAZs explain the development of the dramatic differences in the distribution of ipsilateral and contralateral RGCs.

We have changed the title to avoid implying a causal relationship between clustering and eye-specific segregation. Our key findings in Figures 4C and 4D demonstrate effect sizes >2.0 with high statistical power (Supplemental Table S2). While the absolute magnitude of differences is modest (5-7%), these high effect sizes combined with low inter-animal variability demonstrate consistent, reproducible biological phenomena. During development, small differences during critical periods can have profound downstream consequences for synaptic refinement outcomes.

We acknowledge that significance in Figure 4 arises due to low variance between biological replicates rather than large mean differences. We have revised the text to describe these as "slight" differences and that "WT mice show a tendency toward forming more synapses near mAZ inputs," reflecting appropriate caution in our interpretation while maintaining the statistical robustness of our findings.

(2) Sample size. Performing a large number of significance tests and comparing pvalues is not hypothesis testing and is not descriptive science. At best, with large sample sizes and controls for multiple tests, this approach could be considered exploratory. With n=3 for each group, many comparisons of many derived measures, among many groups, and no control for multiple testing, this approach constitutes a random result generator.

The authors argue that n=3 is a large sample size for the type of high resolution / large volume data being used. It is true that many electron microscopy studies with n=1 are used to reveal the patterns of organization that are possible within an individual. However, such studies cannot control individual variation and are, therefore, not appropriate for identifying subtle differences between groups.

In response to previous critiques along these lines, the authors argue they have dealt with this issue by limiting their analysis to within-individual paired comparisons. There are several problems with their thinking in this approach. The main problem is that they did not change the logic of their arguments, only which direction they pointed the t-tests. Instead of claiming that two groups are different because p < 0.05, they say that two groups are different because one produced p < 0.05 and the other produced p > 0.05. These arguments are not statistically valid or biologically meaningful.

We have implemented rigorous statistical controls, applying false discovery rate (FDR) correction using the Benjamini-Hochberg method (α = 0.05) within each experimental condition (age × genotype combination). This correction strategy treats each condition as addressing a distinct experimental question: “What synaptic properties differ between left eye and right eye inputs in this specific developmental stage and genotype?” The approach appropriately controls for multiple testing while preserving power to detect biologically meaningful differences. We applied FDR correction separately to the ~20-34 measurements (varying by age and genotype) within each of the six experimental conditions, resulting in condition-specific adjusted p-values reported in updated Supplemental Table S2. This correction confirmed the robustness of our key findings. We do not base conclusions solely on comparing p-values across conditions. Our interpretations focus on effect sizes, confidence intervals, and consistent patterns within each condition, with statistical significance providing supporting evidence rather than the primary basis for biological conclusions.

To the best of my understanding, the results are consistent with the following model:

RGCs form mAZs at large boutons (known)

About a quarter of week-one RGC boutons are mAZs (new observation)

Vesicle clustering is proportional to active zone number (~new observation)

RGC synapse density increases during the first post-week (known)

Blocking activity reduces synapse density (known)

Contralateral eye RGCs for more and larger synapses in the lateral dLGN (known)

While mAZ formation is known in adult and juvenile dLGN, the formation of mAZ boutons during eye-specific competition represents new information with important functional implications. Synapses with multiple release sites should be stronger than single-active-zone synapses, suggesting a structural correlate for competitive advantage during refinement.

We demonstrate distinct developmental patterns for sAZ versus mAZ contacts during the first postnatal week. Multi-active zone density favors the dominant eye, while single active-zone synapse density from the competing eye increases from P2-P4 to match dominant-eye levels. This reveals that newly formed synapses from the competing eye predominantly contain single release sites, marking P4-P8 as a critical window for understanding molecular mechanisms driving synaptic elimination.

Our results show that altered retinal activity patterns (β2KO mice) reduce synapse density during eye-specific competition. We relied on β2 knockout mice, which retain retinal waves and spontaneous spike activity but with disrupted patterns and output levels compared to controls. We make no claims about complete activity blockade. Previous studies using different activity manipulations (epibatidine, TTX) have examined terminal morphology, but effects on synapse density during competition remain largely unknown. Achieving complete retinal activity blockade is technically challenging, making it of interest to revisit the role of activity using more precise manipulations to control spike output and relative timing.

With n=3 and effect sizes smaller than 1 standard deviation, a statistically significant result is about as likely to be a false positive as a true positive.

A true-positive statistically significant result does is not evidence of a meaningful deviation from a biological model.

Our conclusions are based on results with effect sizes substantially larger than 1. Key findings demonstrate effect sizes exceeding 2.0. These large effect sizes, combined with rigorous FDR correction and low inter-animal variability, provide evidence against false positive results. During critical developmental periods, consistent structural differences, even those modest in absolute magnitude, can reflect important regulatory mechanisms that influence refinement outcomes. All statistical results, effect sizes, and power analyses are reported in Supplementary Tables S2, with confidence intervals in Supplementary Table S3. We have revised the text in several places where small differences are presented to reflect appropriate caution in our interpretation.

Providing plots that show the number of active zones present in boutons across these various conditions is useful. However, I could find no compelling deviation from the above default predictions that would influence how I see the role of mAZs in activity dependent eye-specific segregation.

Below are critiques of most of the claims of the manuscript.

Claim (abstract): individual retinogeniculate boutons begin forming multiple nearby presynaptic active zones during the first postnatal week.

Confirmed by data.

Claim (abstract): the dominant-eye forms more numerous mAZ contacts,

Misleading: The dominant-eye (by definition) forms more contacts than the nondominant eye. That includes mAZ.

While the dominant eye forms more total contacts, the pattern depends critically on contact type and developmental stage. The dominant eye forms more mAZ contacts across all ages (Figures 2 and S1). However, for sAZ contacts, the two eyes form similar numbers at P4, with the non-dominant eye showing increased sAZ formation during this critical period. This differential pattern by synapse type represents an important aspect of how synaptic competition unfolds structurally.

Claim (abstract): At the height of competition, the non-dominant-eye projection adds many single active zone (sAZ) synapses

Weak: While the individual observation is strong, it is a surprising deviation based on a single n=3 experiment in a study that performed twelve such experiments (six ages, mutant/wildtype, sAZ/mAZ)

The difference in eye-specific sAZ formation at P2 and P8 had effect sizes of ~5.3 and ~2.7 respectively (after FDR correction the difference was still significant at P2 and trending at P8). At P4, no effect was observed by paired T-test and the 5/95% confidence intervals ranged from -0.021-0.008 synapses/m3. The consistency of this pattern across P2 and P8, combined with the large effect sizes, supports the reliability of this developmental finding. We report all effect sizes and power test analyses in Supplemental Table S2, and confidence intervals in Supplemental Table S3.

Claim (abstract): Together, these findings reveal eye-specific differences in release site addition during synaptic competition in circuits essential for visual perception and behavior.

False: This claim is unambiguously false. The above findings, even if true, do not argue for any functional significance to active zone clustering.

Our phrasing “circuits essential for visual perception and behavior” referred to the general importance of binocular organization in the retinogeniculate system for visual processing and we did not intend to claim direct functional significance of our structural data. For clarity we have deleted the latter part of this sentence. In lines 35-37, the abstract now reads “Together, these findings reveal eye-specific differences in release site addition that correlate with axonal refinement outcomes during retinogeniculate refinement.”

Claim (line 84): "At the peak of synaptic competition midway through the first postnatal week, the non-dominant-eye formed numerous sAZ inputs, equalizing the global synapse density between the two eyes"

Weak: At one of twelve measures (age, bouton type, genotype) performed with 3 mice each, one density measure was about twice as high as expected.

The difference in eye-specific sAZ formation at P2 and P8 had effect sizes of ~5.3 and ~2.7 respectively (after FDR correction the difference was still significant at P2 and trending at P8). At P4, no effect was observed by paired T-test and the 5/95% confidence intervals ranged from -0.021-0.008 synapses/m3. The consistency of this pattern across P2 and P8, combined with the large effect sizes, supports the reliability of this developmental finding. We report all effect sizes and power test analyses in Supplemental Table S2, and confidence intervals in Supplemental Table S3.

Claim (line 172): "In WT mice, both mAZ (Fig. 3A, left) and sAZ (Fig. 3B, left) inputs showed significant eye-specific volume differences at each age."

Questionable: There appears to be a trend, but the size and consistency is unclear.

Claim (line 175): "the median VGluT2 cluster volume in dominant-eye mAZ inputs was 3.72 fold larger than that of non-dominant-eye inputs (Fig. 3A, left)."

Cherry picking. Twelve differences were measured with an n of 3, 3 each time. The biggest difference of the group was cited. No analysis is provided for the range of uncertainty about this measure (2.5 standard deviations) as an individual sample or as one of twelve comparisons.

Claim (line 174): "In the middle of eye-specific competition at P4 in WT mice, the median VGluT2 cluster volume in dominant-eye mAZ inputs was 3.72 fold larger than that of non-dominant-eye inputs (Fig. 3A, left). In contrast, β2KO mice showed a smaller 1.1 fold difference at the same age (Fig. 3A, right panel). For sAZ synapses at P4, the magnitudes of eye-specific differences in VGluT2 volume were smaller: 1.35-fold in WT (Fig. 3B, left) and 0.41-fold in β2KO mice (Fig. 3B, right). Thus, both mAZ and sAZ input size favors the dominant eye, with larger eye-specific differences seen in WT mice (see Table S3)."

No way to judge the reliability of the analysis and trivial conclusion: To analyze effect size the authors choose the median value of three measures (whatever the middle value is). They then make four comparisons at the time point where they observed the biggest difference in favor of their hypothesis. There is no way to determine how much we should trust these numbers besides spending time with the mislabeled scatter plots. The authors then claim that this analysis provides evidence that there is a difference in vGluT2 cluster volume between dominant and non-dominant RGCs and that that difference is activity dependent. The conclusion that dominant axons have bigger boutons and that mutants that lack the property that would drive segregation would show less of a difference is very consistent with the literature. Moreover, there is no context provided about what 1.35 or 1.1 fold difference means for the biology of the system.

We focused on P4 for biological reasons rather than post-hoc selection. P4 represents the established peak of synaptic competition when eye-specific synapse densities are globally equivalent. This is a timepoint consistently highlighted throughout our manuscript and supported by previous literature. We have modified our presentation from fold changes to measured eye-specific differences in volume (mean ± standard error) and added confidence intervals in Supplemental Table S3. The effect sizes for eye-specific differences in VGluT2 volume at P4 are robust: ~2.3 and ~1.5 for mAZ and sAZ measurements in WT mice, and ~2.5 and ~1.8 in β2KO mice, with all analyses well-powered (Supplemental Table S2).

We were unable to identify any mislabeled scatter plots and believe all figures are correctly labeled. While dominant-eye advantage in bouton size is consistent with previous literature, our study provides the first detailed analysis of how this develops specifically during the critical period of competition, with distinct patterns for single versus multi-active zone contacts. Our data show that dominant-eye inputs have larger vesicle pools that scale with active zone number. While this suggests enhanced transmission capacity, we make no direct physiological claims based on structural data alone.

Claim (189): "This shows that vesicle docking at release sites favors the dominant-eye as we previously reported but is similar for like eye type inputs regardless of AZ number."

Contradicts core claim of manuscript: Consistent with previous literature, there is an activity dependent relative increase in vGlut2 clustering of dominant eye RGCs. The new information is that that activity dependence is more or less the same in sAZ and mAZ. The only plausible alternative is that vGlut2 scaling only increases in mAZ which would be consistent with the claims of their paper. That is not what they found. To the extent that the analysis presented in this manuscript tests a hypothesis, this is it. The claim of the title has been refuted by figure 3.

We report the volume of docked vesicle signal (VGluT2) nearby each active zone, finding this is greater for dominant-eye synapses. Within each eye-specific synapse population, vesicle signal per active zone is similar regardless of whether these are part of single- or multi-active zone contacts. This is consistent with a modular program of active zone assembly and maintenance: core molecular programs facilitate docking at each AZ similarly regardless of how many AZs are nearby.

This finding does not contradict our main conclusions but rather provides insight into how synaptic advantages are structured. The dominant eye's advantage may arise in part from forming more multi-AZ contacts (which have proportionally more docked vesicles) rather than from enhanced vesicle loading per individual active zone. This organization may reflect how developmental competition operates through contact number and active zone addition rather than fundamental changes to individual release site properties.

We have changed the title to be descriptive rather than mechanistic.

Claim (line 235): "For the non-dominant eye projection, however, clustered mAZ inputs outnumbered clustered sAZ inputs at P4 (Fig. 4C, bottom left panel), the age when this eye adds sAZ synapses (Fig. 2C)."

Misleading: The overwhelming trend across 24 comparisons is that the sAZ clustering looks like mAZ clustering. That is the objective and unambiguous result. Among these 24 underpowered tests (n=3), there were a few p-values < 0.05. The authors base their interpretation of cell behavior on these crossings.

In Figures 4C and 4D we report significant results with high effect sizes (effect sizes all greater than 2; see Supplemental Table S2). The mean differences are modest (5-7%) and significance arises due to low variance between biological replicates. We acknowledge that clustering patterns are generally similar between mAZ and sAZ inputs across most conditions. We have revised the text to describe these as “slight” differences and that “WT mice show a tendency toward forming more synapses near mAZ inputs”, reflecting appropriate caution in our interpretation while noting the statistical consistency of these patterns.

Claim (line 328): "The failure to add synapses reduced synaptic clustering and more inputs formed in isolation in the mutants compared to controls."

Trivially true: Density was lower in mutant.

We have rewritten the sentence for clarity: “The failure to add synapses could explain the observation that synaptic clustering was reduced and more inputs formed in isolation in the mutants compared to controls.”

Claim (line 332): "While our findings support a role for spontaneous retinal activity in presynaptic release site addition and clustering..."

Not meaningfully supported by evidence: I could not find meaningful differences between WT and mutant beside the already known dramatic difference in synapse density.

We have changed the sentence to avoid overinterpreting the results. The new sentence in lines 415-417 reads: “While our results highlight developmental changes in presynaptic release site addition and clustering, activity-dependent postsynaptic mechanisms also influence input refinement at later stages.”

Reviewer #2 (Public review):

Summary:

In this manuscript, Zhang and Speer examine changes in the spatial organization of synaptic proteins during eye specific segregation, a developmental period when axons from the two eyes initially mingle and gradually segregate into eye-specific regions of the dorsal lateral geniculate. The authors use STORM microscopy and immunostain presynaptic (VGluT2, Bassoon) and postsynaptic (Homer) proteins to identify synaptic release sites. Activity-dependent changes of this spatial organization are identified by comparing the β2KO mice to WT mice. They describe two types of synapses based on Bassoon clustering: the multiple active zone (mAZ) synapse and single active zone (sAZ) synapse. In this revision, the authors have added EM data to support the idea that mAZ synapses represent boutons with multiple release sites. They have also reanalyzed their data set with different statistical approaches.

Strengths:

The data presented is of good quality and provides an unprecedented view at high resolution of the presynaptic components of the retinogeniculate synapse during active developmental remodeling. This approach offers an advance to the previous mouse EM studies of this synapse because of the CTB label allows identification of the eye from which the presynaptic terminal arises.

Weaknesses:

While the interpretation of this data set is much more grounded in this second revised submission, some of the authors' conclusions/statements still lack convincing supporting evidence. In particular, the data does not support the title: "Eye-specific active zone clustering underlies synaptic competition in the developing visual system". The data show that there are fewer synapses made for both contra- and ipsi- inputs in the β2KO mice-- this fact alone can account for the differences in clustering. There is no evidence linking clustering to synaptic competition. Moreover, the findings of differences in AZ# or distance between AZs that the authors report are quite small and it is not clear whether they are functionally meaningful.

We thank the reviewer for their helpful suggestions that improved the manuscript in this revision. We have changed the title to remove the reference to “clustering” and to avoid implying any causal relationships. The new title is descriptive: “Eye-specific differences in active zone addition during synaptic competition in the developing visual system”.

To further address the reviewers comments, we have removed the remaining references to activity-dependent effects on synaptic development (line 36, line 96, line 415). We have also modified the text in lines 411-413 to state that “The failure to add synapses could explain the observation that synaptic clustering was reduced and more inputs formed in isolation in the mutants compared to controls.”

We have also updated our presentation of results for Figure 4 to ensure that we do not causally link clustering to synaptic competition. In Figures 4C and 4D we report significant results with high effect sizes (effect sizes all greater than 2; see Supplemental Table S2). The mean differences are modest (5-7%) and significance arises due to low variance between biological replicates. We acknowledge that clustering patterns are generally similar between mAZ and sAZ inputs across most conditions. We have revised the text to describe these as “slight” differences and that “WT mice show a tendency toward forming more synapses near mAZ inputs”, reflecting appropriate caution in our interpretation while noting the statistical consistency of these patterns.

Reviewer #3 (Public review):

This study is a follow-up to a recent study of synaptic development based on a powerful data set that combines anterograde labeling, immunofluorescence labeling of synaptic proteins, and STORM imaging (Cell Reports, 2023). Specifically, they use anti-Vglut2 label to determine the size of the presynaptic structure (which they describe as the vesicle pool size), anti-Bassoon to label active zones with the resolution to count them, and anti-Homer to identify postsynaptic densities. Their previous study compared the detailed synaptic structure across the development of synapses made with contraprojecting vs. ipsi-projecting RGCs and compared this developmental profile with a mouse model with reduced retinal waves. In this study, they produce a new detailed analysis on the same data set in which they classify synapses into "multi-active zone" vs. "single-active zone" synapses and assess the number and spacing of these synapses. The authors use measurements to make conclusions about the role of retinal waves in the generation of same-eye synaptic clusters. The authors interpret these results as providing insight into how neural activity drives synapse maturation, the strength of their conclusions is not directly tested by their analysis.

Strengths:

This is a fantastic data set for describing the structural details of synapse development in a part of the brain undergoing activity-dependent synaptic rearrangements. The fact that they can differentiate the eye of origin is what makes this data set unique over previous structural work. The addition of example images from the EM dataset provides confidence in their categorization scheme.

Weaknesses:

Though the descriptions of single vs multi-active zone synapses are important and represent a significant advance, the authors continue to make unsupported conclusions regarding the biological processes driving these changes. Although this revision includes additional information about the populations tested and the tests conducted, the authors do not address the issue raised by previous reviews. Specifically, they provide no assessment of what effect size represents a biologically meaningful result. For example, a more appropriate title is "The distribution of eye-specific single vs multiactive zone is altered in mice with reduced spontaneous activity" rather than concluding that this difference in clustering is somehow related to synaptic competition. Of course, the authors are free to speculate, but many of the conclusions of the paper are not supported by their results.

We appreciate the reviewer’s helpful critique. We have changed the title to be descriptive and avoid implying causal relationships.

We have applied false discovery rate (FDR) correction using the Benjamini-Hochberg method with α = 0.05 within each experimental condition (age × genotype combination). The FDR correction treats each condition as addressing a distinct experimental question: 'What synaptic properties differ between left eye and right eye inputs in this specific developmental stage and genotype?'

This correction strategy is appropriate because: 1) we focus our statistical comparisons within each age/genotype; 2) each age-genotype combination represents a separate biological context where different synaptic properties between eye-of-origin may be relevant; and 3) this approach controls for multiple testing within each experimental question while maintaining statistical power to detect meaningful biological differences.

We applied FDR correction separately to the ~20-34 measurements (varying with age and genotype) within each of the six experimental conditions (P2-WT, P2-ß2, P4-WT, P4-ß2, P8-WT, P8-ß2), resulting in condition-specific adjusted p-values. These are reported in the updated Supplemental Table S2. Figures have been also been updated to reflect the FDR-adjusted values. Selected between-genotype comparisons are presented descriptively using 5/95% confidence intervals. This correction confirmed the robustness of our key findings.

With regard to the biological significance of effect sizes, our key findings demonstrate effect sizes >2.0, indicating robust effects. During critical developmental periods, consistent structural differences, even those modest in absolute magnitude, can reflect important regulatory mechanisms that influence refinement outcomes. The differences in synaptic organization we observe occur during the first postnatal week when eyespecific competition is active, suggesting these patterns may be relevant to understanding how structural advantages emerge during synaptic refinement.

Reviewer #1 (Recommendations for the authors):

I have tried to understand the analysis and biology of this manuscript as best I can. I believe the analytical approach taken is not reliable and I have explained why in my public comments. I don't believe this manuscript is unique in taking this approach. I have recently published a paper on how common this approach is and why it doesn't work. I don't want to give the impression that the problem with the analysis was that it was not computationally sophisticated enough or that you did not jump through a specific statistical hoop. If I strip out the arguments that depend on misinterpretations of p-values and -instead- look at the scatterplots, I come up with a very different view of the data than what is described in the paper.

The information in the plots could be translated into a rigorous statistical analysis of estimated differences between groups given the uncertainties of the experimental design. I don't really think that analysis would be useful. I think it would have been enough to publish the plots and report your estimates of the number of active zones in RGCs during development. I don't see evidence of an additional effect.

We appreciate the reviewer’s helpful comments throughout the review process. Mean active zone numbers per mAZ contact are presented in Figure S2D/E. We look forward to further technical and computational advances that will help us increase our data acquisition throughput and sample sizes when designing future studies.

Reviewer #2 (Recommendations for the authors):

The authors should modify the title and other text to be more consistent with the data. There is no evidence that active zone clustering has any direct relationship to synaptic competition.

We appreciate the reviewer’s helpful suggestions to ensure appropriate language around causal effects. We have modified the title to accurately reflect the results: "Eyespecific differences in active zone addition during synaptic competition in the developing visual system." We have revised the text in the abstract, introduction, and results section for Figures 4 to be consistent with the data and not imply causality of synapse clustering on segregation phenotypes.

Reviewer #3 (Recommendations for the authors):

Change the title.

We appreciate the reviewer’s feedback throughout the review process. We have modified the title to accurately reflect the results: "Eye-specific differences in active zone addition during synaptic competition in the developing visual system."

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Spatially clustered synaptic inputs enable local dendritic computations important for learning, memory, and sensory processing. In the mammalian visual system, individual retinal ganglion cell (RGC) axons form clustered terminal boutons containing multiple active zones onto relay cell dendrites in the dorsal lateral geniculate nucleus (dLGN). This mature architecture arises through the addition of release sites, which strengthens selected afferents while weaker inputs are pruned. Following eye-opening, spontaneous activity and visual experience promote synaptic refinement and bouton clustering after binocular inputs have segregated. However, anatomical changes in release site addition and spatial patterning during earlier stages of eye-specific competition are not well understood. To investigate this, we examined the spatial organization of eye-specific active zones in wild type mice and a mutant line with disrupted cholinergic retinal waves. Using volumetric super-resolution single-molecule localization microscopy and electron microscopy, we found that individual retinogeniculate boutons begin forming multiple nearby presynaptic active zones during the first postnatal week. Both eyes generate these “multi-active-zone” (mAZ) inputs throughout refinement, but the dominant-eye forms more numerous mAZ contacts, each with more active zones and larger vesicle pools. At the height of competition, the non-dominant-eye projection adds many single active zone (sAZ) synapses. Mutants with abnormal cholinergic retinal waves still form mAZ inputs, but develop fewer sAZ synapses and show reduced synapse clustering in projections from both eyes. Together, these findings reveal activity-dependent, eye-specific differences in release site addition during synaptic competition in circuits essential for visual perception and behavior.

Revised version: Spatially clustered synaptic inputs enable local dendritic computations important for learning, memory, and sensory processing. In the mammalian visual system, individual retinal ganglion cell (RGC) axons form clustered terminal boutons containing multiple active zones onto relay cell dendrites in the dorsal lateral geniculate nucleus (dLGN). This mature architecture arises through the addition of release sites, which strengthens selected afferents while weaker inputs are pruned. Following eye-opening, spontaneous activity and visual experience promote synaptic refinement and bouton clustering after binocular inputs have segregated. However, anatomical changes in release site addition and spatial patterning during earlier stages of eye-specific competition are not well understood. To investigate this, we examined the spatial organization of eye-specific active zones in wild type mice and a mutant line with disrupted cholinergic retinal waves. Using volumetric super-resolution single-molecule localization microscopy and electron microscopy, we found that individual retinogeniculate boutons begin forming multiple nearby presynaptic active zones during the first postnatal week. Both eyes generate these “multi-active-zone” (mAZ) inputs throughout refinement, but the dominant-eye forms more numerous mAZ contacts, each with more active zones and larger vesicle pools. At the height of competition (postnatal day 4), the non-dominant-eye projection adds many single active zone (sAZ) synapses. Mutants with abnormal cholinergic retinal waves still form mAZ inputs, but develop fewer synapses overall and show reduced synaptic clustering in projections from both eyes. Together, these findings reveal eye-specific differences in release site addition that correlate with axonal refinement outcomes during retinogeniculate refinement.

Introduction

A key mechanism for neuronal signal processing is the formation of spatially clustered synapses that facilitate local computations within individual dendrites 1–3. Through biophysical signal integration mechanisms, neighboring synaptic inputs play critical computational roles in learning, memory, and sensory processing underlying cognition and behavior 4–7. During circuit development, both spontaneous and sensory-driven neural activity regulate synaptic clustering, stabilizing some synapses and eliminating others to establish mature connectivity patterns 6–8.

The refinement of retinal inputs to the dorsal lateral geniculate nucleus (dLGN) of the thalamus is a model example of activity-dependent synaptic clustering 9. Over development, boutons within the terminal arbors of individual retinogeniculate axons cluster progressively 10–13, creating multiple synaptic contacts onto individual postsynaptic targets 14–16. Concurrently, developing RGC boutons add presynaptic release sites such that some mature retinogeniculate terminals contain several dozen active zones (AZs) 16, 17. The formation of clustered boutons with multiple release sites subserves a critical “driver” function of retinogeniculate input 18, 19.

Earlier version: Retinogeniculate bouton clustering depends on visual experience, and dark-rearing after eye-opening reduces clustering within individual axon arbors 12. However, less is known about synaptic spatial relationships and activity-dependent release site refinement during eye-specific competition before eye-opening. During early axon ingrowth, individual RGCs extend sparse side branches in the incorrect eye-specific territory 20–22. These transient branches form synapses that contain fewer presynaptic vesicles than those made by the same axon in the correct eye-specific domain 23.

Revised version: Retinogeniculate bouton clustering depends on visual experience, and dark-rearing after eye-opening reduces clustering within individual axon arbors 12. However, less is known about synaptic spatial relationships and activity-dependent release site refinement during eye-specific competition before eye-opening. During early axon ingrowth, individual RGCs extend sparse side branches in the incorrect eye-specific territory 20–22. These transient branches form synapses that contain fewer presynaptic vesicles than those made by the same axon in the correct eye-specific domain 23. Functional recordings after eye-opening show that while many RGC inputs are pruned, the remaining inputs strengthen by adding release sites 24. This leads to glutamate spillover and cross talk between RGC inputs that increases postsynaptic relay neuron excitability 25. Thus, early eye-specific differences in release site number or spacing may contribute to binocular input refinement prior to eye-opening.

Removed in the revised version: Functional recordings after eye-opening show that while many RGC inputs are pruned, the remaining inputs strengthen by adding release sites 24. This leads to glutamate spillover and cross talk between RGC inputs that increases postsynaptic relay neuron excitability 25. Thus, early eye-specific differences in release site number or spacing may contribute to binocular input refinement prior to eye-opening.

In a previous study from our laboratory, we used volumetric stochastic optical reconstruction microscopy (STORM), anterograde tract tracing, and immunohistochemical labeling of synaptic proteins to show that dominant-eye synapses contain more total vesicles and have more vesicles near the AZ (putative docking) compared with non-dominant-eye synapses 26. However, we did not examine whether eye-specific inputs differ in AZ number or AZ spatial proximity (clustering) during synaptic competition. To address this gap, we reanalyzed our published dataset, focusing on eye-specific AZ spatial relationships. We found that projections from each eye form a subset of retinogeniculate inputs in which several (2-4) AZs share a common cluster of presynaptic vesicles. We refer to these synapses as multi-active-zone (mAZ) inputs. All other retinogeniculate inputs contained one single active zone and its associated vesicle cluster, and we term these single-active-zone (sAZ) synapses.

Earlier version: During eye-specific synaptic competition, the dominant-eye projection formed more mAZ inputs, each with more AZs and a larger presynaptic vesicle pool compared to the non-dominant eye projection. Similarly, the dominant-eye had higher vesicle signal at sAZ inputs. At the peak of synaptic competition midway through the first postnatal week, the non-dominant-eye formed numerous sAZ inputs, equalizing the global synapse density between the two eyes. These eye-specific AZ patterns were disrupted in a mutant mouse line with abnormal stage II cholinergic retinal waves and retinogeniculate segregation defects. Our results demonstrate that eye-specific presynaptic release site clustering arises before eye-opening and is impacted by defects in spontaneous retinal activity.

Revised version: During eye-specific synaptic competition, the dominant-eye projection formed more mAZ inputs, each with more AZs and a larger presynaptic vesicle pool compared to the non-dominant eye projection. Similarly, the dominant-eye had higher vesicle signal at sAZ inputs. At the peak of synaptic competition midway through the first postnatal week (postnatal day 4), the non-dominant-eye formed numerous sAZ inputs, equalizing the global synapse density between the two eyes. These eye-specific AZ patterns were disrupted in a mutant mouse line with abnormal stage II cholinergic retinal waves and retinogeniculate segregation defects.

Results

Retinogeniculate inputs form multiple active zones during eye-specific competition

To investigate active zone refinement during eye-specific segregation, we reanalyzed a volumetric super-resolution imaging dataset previously published by our laboratory 26. We used volumetric STochastic Optical Reconstruction Microscopy (STORM) 27 to image the dLGNs of wild-type (WT) mice at three postnatal ages (P2, P4, and P8) (Fig. 1A). We labeled eye-specific inputs by monocular injection of Alexa Fluor-conjugated cholera toxin subunit B tracer (CTB) together with immunostaining for presynaptic Bassoon, postsynaptic Homer1, and presynaptic vesicular glutamate transporter 2 (VGluT2) proteins (Fig. 1B). We collected separate image volumes (∼45K μm3 each) from three biological replicates at each age. To assess the impact of spontaneous retinal activity on synaptic development across the same time period, we performed identical experiments in β2-knockout (β2KO) mice lacking the beta 2 subunit of the nicotinic acetylcholine receptor, a mutation that disrupts spontaneous cholinergic retinal wave activity, eye-specific segregation, and retinogeniculate synapse development 22, 26, 28–36 (Fig. 1A). Because eye-specific segregation is incomplete until ∼P8, we limited our re-analysis to the future contralateral eye-specific region of the dLGN, which is reliably identified across all stages of postnatal development (Fig. 1A, see also Materials and Methods).

Retinogeniculate boutons form multiple active zones during eye-specific competition

(A) Experimental design. CTB-Alexa 488 was injected into the right eye of wild-type and β2KO mice. One day after the treatment, tissue was collected from the left dLGN at P2, P4, and P8. Red squares indicate the STORM imaging regions that were analyzed. (B) Representative examples of individual sAZ and mAZ inputs, with corresponding active zone counts ranging from one to three. Upper panels show Z-projections of inputs and lower panels show the corresponding 3D volume. Arrowheads point to individual Bassoon clusters (active zones) paired with postsynaptic Homer1 labels within each input. All examples are from a WT P8 sample. (C) Electron micrographs of mAZ retinogeniculate inputs in a P8 ET33Cre::ROSA26LSL-Matrix-dAPEX2 mouse. Darkly-stained dAPEX2(+) mitochondria are present within ipsilaterally-projecting RGC terminals. Arrowheads point to electron-dense material at the postsynaptic density, apposed to individual active zones with clustered presynaptic synaptic vesicles.

By analyzing synapses in the contralateral dLGN from eighteen mice across three ages and two genotypes (Table S1), STORM revealed two classes of retinogeniculate inputs distinguished by active zone (AZ) number (Fig. 1B). We defined each retinogeniculate input as a single contiguous VGluT2 cluster together with all its associated presynaptic (Bassoon) and postsynaptic (Homer1) paired synaptic labels. Using this definition, inputs that had multiple (2-4) Bassoon AZs were classified as multi-active-zone (mAZ) inputs, while those with a single Bassoon AZ were designated single-active-zone (sAZ) inputs (Fig. 1B). Most mAZ inputs contained two AZs (∼70-90%, varying with age, genotype, and eye-of-origin); smaller proportions contained three AZs (∼10-20%) or four or more AZs (<5%) (Table S1).

Each mAZ input could be a single terminal bouton with several AZs, or a cluster of sAZ synapses within separate boutons 10, 11, 15, 16, 37. To address this, we used electron microscopy (EM) to image retinogeniculate terminals in the dLGN at P8. We generated a transgenic line expressing mitochondrial matrix-targeted dimeric dAPEX2 reporter 38 in ipsilaterally-projecting RGCs 39–41, providing unambiguous mitochondrial labeling in ipsiRGC axons. EM images confirmed the presence of individual retinogeniculate boutons with multiple active zones, consistent with our STORM data (Fig. 1C). Previous EM reconstructions of retinogeniculate inputs reported no evidence of RGC bouton convergence at the end of the first postnatal week 11. Together, these results suggest that mAZ inputs in STORM images are single RGC terminals that house several closely spaced release sites. Hereafter, we use the word “synapse” only to refer to each partnered active zone (Bassoon/Homer1 pairs); mAZ inputs contain several synapses, while each sAZ input is one synapse.

Changes in eye-specific input density during synaptic competition

In our previous analysis, we reported global eye-specific synapse densities that reflected the combined mAZ and sAZ inputs. Dominant-eye synapse density was greater than that of the non-dominant eye at P2 and P8, but eye-specific synapse densities were equivalent at P4 during the peak of synaptic competition 26. To determine how these two input classes evolve during competition, we quantified the densities and percentages of mAZ and sAZ inputs in WT and β2KO mice (Fig. 2A). Eye-of-origin for each retinogeniculate input was assigned by colocalizing CTB signal with VGluT2 (Fig. 2B). Binocular CTB control injections showed that anterograde tracing labeled >97% of retinogeniculate synapses at P4 and P8, ensuring accurate eye-specific assignment 26. Within the contralateral eye-specific region of the dLGN, CTB(+) VGluT2 clusters were classified as “dominant-eye” inputs and CTB(-) VGluT2 clusters as “non-dominant eye” inputs (Fig. 2B).

Changes in eye-specific input density during synaptic competition

Earlier version: (A) Representative Z-projection images of mAZ and sAZ inputs across ages and genotypes. Arrowheads point to individual Bassoon/Homer1 cluster pairs indicating release sites. (B) Representative CTB(+) dominant-eye (top panels) and CTB(-) non-dominant-eye (bottom panels) mAZ inputs in a WT P8 sample, showing synaptic (left panels), CTB (middle panels), and merged labels (right panels). Arrowheads point to individual Bassoon/Homer1 paired clusters. (C) Eye-specific mAZ (left) and sAZ (right) input density across development in WT (top panels) and β2KO mice (bottom panels). Black dots represent mean values from separate biological replicates and black lines connect eye-specific measurements within each replicate (N=3 for each age and genotype). Error bars represent group means ± SEMs. Statistical significance between eye-specific measurements was assessed using two-tailed paired T-tests at each age. *p<0.05, **p<0.01.

Revised version: (A) Representative Z-projection images of mAZ and sAZ inputs across ages and genotypes. Arrowheads point to individual Bassoon/Homer1 cluster pairs indicating release sites. (B) Representative CTB(+) dominant-eye (top panels) and CTB(-) non-dominant-eye (bottom panels) mAZ inputs in a WT P8 sample, showing synaptic (left panels), CTB (middle panels), and merged labels (right panels). Arrowheads point to individual Bassoon/Homer1 paired clusters. (C) Eye-specific mAZ (left) and sAZ (right) input density across development in WT (top panels) and β2KO mice (bottom panels). Black dots represent mean values from separate biological replicates and black lines connect eye-specific measurements within each replicate (N=3 for each age and genotype). Error bars represent group means ± SEMs. Statistical significance between eye-specific measurements was assessed for each genotype using two-tailed paired T-tests with Benjamini-Hochberg FDR correction (α = 0.05) at each age. *p(adj)<0.05.

Removed in the revised version: After assigning eye-of-origin to each retinogeniculate synapse, we found that the density of CTB(+) mAZ inputs was significantly higher than CTB(-) mAZ inputs at every age in WT mice (Fig. 2C, top left; Fig. S1, top). β2KO mice also developed eye-specific differences in mAZ input density (Fig. 2C, bottom left; Fig. S1, bottom). For WT sAZ synapses, the dominant-eye had a higher synapse density at P2 and P8. At P4, however, sAZ synapse density was equivalent between the eyes (Fig. 2C, top right).

Earlier version: This pattern resulted from non-dominant eye sAZ synapse addition, which equalized the global input density between the eyes at P4 as we reported previously (5/95% confidence interval,-0.014-0.011 synapses / μm3). In β2KO mice, the non-dominant eye formed fewer sAZ synapses, trending toward an eye-specific difference. (Fig. 2C, bottom right; p<0.05, power=0.62; see Tables S2/S3). Thus, β2KO mice with disrupted retinal activity maintain a higher mAZ input fraction in the dominant-eye projection, but sAZ synapse addition is reduced at the peak of competition.

Revised version: After false discovery rate (FDR) correction for multiple comparisons within each age and genotype (see Quantification and Statistical Analysis; p(adj) values shown in all figures and Table S2), we found that the density of CTB(+) mAZ inputs tended to be higher than CTB(-) mAZ inputs in WT mice (Fig. 2C, top left). Consistent with this trend, the mAZ input fraction (% of all inputs) was significantly higher for CTB(+) dominant-eye inputs at P4 and P8 (Fig. S1, top). β2KO mice also developed differences in mAZ input density favoring the dominant-eye (Fig. 2C, bottom left; Fig. S1, bottom). For WT sAZ synapses, the dominant-eye had a significantly higher synapse density at P2 [p(adj)=0.026, Cohen’s d = 5.31] and trended higher at P8 [p(adj)=0.072, Cohen’s d = 2.73]. At P4, however, sAZ synapse density was equivalent between the eyes (Fig. 2C, top right; Table S2/S3). This pattern resulted from non-dominant eye sAZ synapse addition, which equalized the global input density between the eyes at P4 as we reported previously (5/95% confidence interval, -0.014-0.011 synapses / μm3). In β2KO mice at P4, the non-dominant eye formed fewer sAZ synapses (Fig. 2C, bottom right; see Tables S2/S3). Thus, β2KO mice with disrupted retinal activity maintain a higher mAZ input fraction in the dominant-eye projection, but sAZ synapse addition is reduced at the peak of competition.

Multi⍰ and single⍰active⍰zone inputs from the dominant eye show increased vesicle pool size and vesicle proximity to the active zone

Earlier version: We previously reported a dominant-eye bias in VGluT2 volume when considering all retinogeniculate inputs 26. Here, we assessed presynaptic maturation separately in sAZ and mAZ inputs by measuring their total VGluT2 volume. In WT mice, both mAZ (Fig. 3A, left) and sAZ (Fig. 3B, left) inputs showed significant eye-specific volume differences at each age. In the middle of eye-specific competition at P4 in WT mice, the median VGluT2 cluster volume in dominant-eye mAZ inputs was 3.72 fold larger than that of non-dominant-eye inputs (Fig. 3A, left). In contrast, β2KO mice showed a smaller 1.1 fold difference at the same age (Fig. 3A, right panel). For sAZ synapses at P4, the magnitudes of eye-specific differences in VGluT2 volume were smaller: 1.35-fold in WT (Fig. 3B, left) and 0.41-fold in β2KO mice (Fig. 3B, right). Thus, both mAZ and sAZ input size favors the dominant eye, with larger eye-specific differences seen in WT mice (see Table S3).

Revised version: We previously reported a dominant-eye bias in VGluT2 volume when considering all retinogeniculate inputs 26. Here, we assessed presynaptic maturation separately in sAZ and mAZ inputs by measuring their total VGluT2 volume. In WT mice, both mAZ (Fig. 3A, left) and sAZ (Fig. 3B, left) inputs showed significant eye-specific volume differences in the middle of eye-specific competition at P4. At this age in WT mice, the median VGluT2 cluster volume in dominant-eye mAZ inputs was ∼3.55 ± 1.3 μm3 larger (mean ± S.E.) than that of non-dominant-eye inputs (Fig. 3A, left). In contrast, β2KO mice showed a smaller ∼1.9 ± 1.1 μm3 (mean ± S.E.) volume difference between median eye-specific mAZ inputs at the same age (Fig. 3A, right panel). For sAZ synapses at P4, the magnitudes of eye-specific differences in median VGluT2 volume (mean ± S.E.) were ∼2.1 ± 1.0 μm3 in WT (Fig. 3B, left) and ∼1.5 ± 1.1 in β2KO mice (Fig. 3B, right). Thus, both mAZ and sAZ vesicle pool volumes are larger for the dominant eye, with the largest eye-specific differences seen for mAZ inputs in WT mice (see Table S3).

Dominant-eye inputs show larger vesicle pools that scale with active zone number

Earlier version: (A and B) Violin plots showing the distribution of VGluT2 cluster volume for (A) mAZ and (B) sAZ inputs in WT (filled) and β2KO mice (striped) at each age. The width of each violin plot reflects the relative synapse proportions across the entire grouped data set at each age (N=3 biological replicates) and the maximum width was normalized across all groups. The black dots represent the median value of each biological replicate (N=3), and the black horizontal lines represent the median value of all inputs grouped across replicates. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Statistical significance was determined using a linear mixed model ANOVA with a post-hoc Bonferroni correction. Black asterisks indicate significant eye-specific differences at each age. *p<0.05, **p<0.01. ***p<0.001. (C) Eye-specific VGluT2 signal volume for all inputs separated by number of AZs in WT (left panel) and β2KO mice (right panel) at P4. (D) Average VGluT2 volume per AZ for all inputs separated by number of AZs in WT (left panel) and β2KO mice (right panel) at P4. In panels (C) and (D), error bars indicate group means ± SEMs (N=3 biological replicates for each age and genotype). Black dots represent mean values from separate biological replicates and black lines connect eye-specific measurements within each replicate. Statistical significance between eye-specific measurements was assessed using two-tailed paired T-tests: *p<0.05; **p<0.01.

Revised version: (A and B) Violin plots showing the distribution of VGluT2 cluster volume for (A) mAZ and (B) sAZ inputs in WT (filled) and β2KO mice (striped) at each age. The width of each violin plot reflects the relative synapse proportions across the entire grouped data set at each age (N=3 biological replicates) and the maximum width was normalized across all groups. The black dots represent the median value of each biological replicate (N=3), and the black horizontal lines represent the median value of all inputs grouped across replicates. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Statistical significance was determined using a linear mixed model ANOVA with a post-hoc Bonferroni correction, followed by Benjamini-Hochberg FDR correction (α = 0.05) for multiple comparisons at each age/genotype. Black asterisks indicate significant eye-specific differences at each age. *p(adj)<0.05. (C) Eye-specific VGluT2 signal volume for all inputs separated by number of AZs in WT (left panel) and β2KO mice (right panel) at P4. (D) Average VGluT2 volume per AZ for all inputs separated by number of AZs in WT (left panel) and β2KO mice (right panel) at P4. In panels (C) and (D), error bars indicate group means ± SEMs (N=3 biological replicates for each age and genotype). Black dots represent mean values from separate biological replicates and black lines connect eye-specific measurements within each replicate. Statistical significance between eye-specific measurements was assessed for each genotype using two-tailed paired T-tests with Benjamini-Hochberg FDR correction (α = 0.05): *p(adj)<0.05.

Earlier version: In addition to total vesicle pool volume, we quantified VGluT2 volume within a 70 nm shell around each AZ in mAZ and sAZ inputs, considering this a proxy for docked vesicles (Fig. S2A-C) 26. In WT mice at P4, dominant-eye inputs showed greater per-AZ vesicle volume than non-dominant inputs in both mAZ and sAZ categories (Fig. S2B, left). These eye-specific differences were absent in β2KO mice (Fig. S2B, right).

Revised version: In addition to total vesicle pool volume, we quantified the readily releasable pool by measuring VGluT2 volume within a 70 nm shell around each AZ, considering this a proxy for docked vesicles (Fig. S2A-C) 26. In WT mice at P4, dominant-eye inputs showed greater vesicle volume per AZ than non-dominant inputs, in both mAZ and sAZ terminals (Fig. S2B, left; Table S3). These eye-specific differences were absent in β2KO mice (Fig. S2B, right; Table S3). However, when comparing mAZ and sAZ inputs from the same eye, vesicle volume per AZ was similar across all ages and genotypes (Fig. S2A-C; Table S2). This confirms our previous finding that vesicle docking favors the dominant-eye 26 and shows that AZs formed by a single eye have similar docking levels in both their mAZ and sAZ terminals.

Removed in the revised version: Moreover, at all ages and in both genotypes, vesicle volume within the 70 nm shell did not differ between mAZ and sAZ inputs (Fig. S2A-C). This shows that vesicle docking at release sites favors the dominant-eye as we previously reported 26, but is similar for like-eye type inputs regardless of AZ number.

Vesicle pool size scales with active zone number

Because mAZ inputs showed greater total VGluT2 volume than sAZ synapses, yet exhibited comparable vesicle docking, the disparity could reflect a scaling effect of vesicle pool size with increased AZ number. To evaluate how presynaptic vesicle pool volume scales with AZ number, we compared the Bassoon cluster number to VGluT2 volume for every retinogeniculate input. In both WT (Fig. S2D) and β2KO mice (Fig. S2E), mAZ inputs contained an average of 2-3 Bassoon clusters (separate AZs) in the first postnatal week. In WT mice, CTB(+) dominant-eye mAZ inputs contained more AZs than CTB(-) non-dominant-eye mAZ inputs at P4 and P8 (Fig. S2D). This maturation was delayed until P8 in β2KO mice (Fig. S2E). For CTB(+) dominant-eye inputs in both genotypes at P4, vesicle pool volume correlated positively with AZ number (Fig. 3C; Pearson correlation coefficients: WT [0.99] and β2KO [0.95]). A similar, but weaker correlation was observed for CTB(-) non-dominant-eye inputs (Pearson correlation coefficients: WT [0.97] and β2KO [0.90]. Dividing the total presynaptic VGluT2 volume by the AZ number revealed a consistent vesicle volume per AZ for both sAZ and mAZ inputs (Fig. 3D/S3, Table S3). Collectively, these results indicate that presynaptic vesicle pool volume scales with AZ number for each eye-specific input, while a dominant-eye bias persists throughout development.

Earlier version: Eye-specific synapse clustering before eye-opening

Revised version: Synapse clustering before eye-opening

Eye-specific competition is thought to involve stabilization of co-active, neighboring inputs from the same eye and elimination of out of sync inputs from the opposite eye 42–44. Because axon refinement depends on the relative strength of neurotransmission between competing inputs 39, 45–51, mAZ inputs with multiple release sites could help stabilize nearby like-eye inputs 48, 52–54.

To quantify synaptic clustering patterns, we measured the distance from every eye-specific sAZ synapse to all other sAZ and mAZ inputs within each image field. sAZ synapses were often found nearby other inputs from the same eye (Fig. 4A). To define clustering, we searched volumetrically around each mAZ input and measured the fraction of sAZ synapses that were nearby at increasing distances (Fig. S4A/B). We then compared the observed percentages with datasets in which sAZ synapse positions were randomly shuffled within each neuropil volume. The largest differences occurred within a 1-2 μm search distance (Fig. S4A/B), and so we chose a 1.5 μm cutoff from the edge of each mAZ input to designate it “clustered” if at least one neighboring sAZ synapse lay within this distance and “isolated” if none did (Fig. 4B). No significant clustering was detected when sAZ and mAZ inputs originated from opposite eyes (Fig. S4C/D).

Eye-specific synapse clustering before eye-opening

Earlier version: (A) Representative mAZ (left panels) and sAZ (right panels) inputs in a WT P8 sample with nearby sAZ synapses (arrowheads) clustered within 1.5 μm (dashed yellow ring). Arrows point to the centered mAZ or sAZ inputs. (B) Ratio of clustered and isolated mAZ and sAZ inputs for CTB(+) (upper panels) and CTB(-) (lower panels) inputs in WT and β2KO mice at P4. (C) Comparison of the clustered input ratio between mAZ and sAZ inputs across different ages, genotypes, and eyes of origin. (D) Comparison of the average number of nearby sAZ synapses for clustered mAZ and sAZ inputs across different ages, genotypes, and eyes of origin. In panels B-D, black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate (N=3 for each age and genotype). Error bars represent group means ± SEMs. Two-tailed paired T-tests were used to test statistical significance between mAZ and sAZ inputs at each age. *p<0.05; **p<0.01.

Revised version: (A) Representative mAZ (left panels) and sAZ (right panels) inputs in a WT P8 sample with nearby sAZ synapses (arrowheads) clustered within 1.5 μm (dashed yellow ring). Arrows point to the centered mAZ or sAZ inputs. (B) Ratio of clustered and isolated mAZ and sAZ inputs for CTB(+) (upper panels) and CTB(-) (lower panels) inputs in WT and β2KO mice at P4. (C) Comparison of the clustered input ratio between mAZ and sAZ inputs across different ages, genotypes, and eyes of origin. (D) Comparison of the average number of nearby sAZ synapses for clustered mAZ and sAZ inputs across different ages, genotypes, and eyes of origin. In panels B-D, black dots represent mean values from separate biological replicates and black lines connect measurements within each replicate (N=3 for each age and genotype). Error bars represent group means ± SEMs. For each genotype, two-tailed paired T-tests with Benjamini-Hochberg FDR correction (α = 0.05) were used to test statistical significance between mAZ and sAZ inputs at each age. *p(adj)<0.05.

Earlier version: At the peak of competition (P4) in WT mice, more than 65% of both mAZ and sAZ inputs were clustered for both eyes, while this proportion fell to ∼50% in β2KO mice (Fig. 4B). In the WT dominant-eye projection, the fractions of clustered mAZ and sAZ inputs were similar at each age (Fig. 4C, top left panel). For the non-dominant eye projection, however, clustered mAZ inputs outnumbered clustered sAZ inputs at P4 (Fig. 4C, bottom left panel), the age when this eye adds sAZ synapses (Fig. 2C). β2KO mice showed no difference between mAZ and sAZ clustering at any age (Fig. 4C, right panels). Additionally, in WT mice at P4 and P8, clustered mAZ inputs from both eyes had more neighboring sAZ synapses than did clustered sAZ synapses (Fig. 4D, left); this enrichment was absent in β2KO mice (Fig. 4D, right). Thus, while most retinogeniculate synapses lie within 1.5 μm of another like-eye input, WT mice show a relative increase in neighboring synapses near mAZ inputs during synaptic competition.

Revised version: At the peak of competition in WT mice (P4), more than 65% of both mAZ and sAZ inputs were clustered for both eyes, while this proportion fell to ∼50% in β2KO mice (Fig. 4B). In the WT dominant-eye projection, the fractions of clustered mAZ and sAZ inputs were similar at each age (Fig. 4C, top left panel). For the non-dominant eye projection, however, there were slightly more clustered mAZ inputs compared to clustered sAZ inputs at P4 (Fig. 4C, bottom left panel), the age when this eye adds sAZ synapses (Fig. 2C). β2KO mice showed no difference between mAZ and sAZ clustering at any age (Fig. 4C, right panels). Additionally, in WT mice at P4 and P8, clustered mAZ inputs from both eyes had marginally more neighboring sAZ synapses than did clustered sAZ synapses (Fig. 4D, left); this enrichment was absent in β2KO mice (Fig. 4D, right). Thus, while most retinogeniculate synapses lie within 1.5 μm of another like-eye input, WT mice show a tendency toward forming more synapses near mAZ inputs during synaptic competition.

Clustered mAZ inputs in P4 WT mice were also closer together than isolated mAZ inputs. Clustered mAZ inputs formed by the dominant-eye were ∼32% closer to the nearest like-eye clustered mAZ input compared to isolated mAZ inputs (Fig. 5A, left panel). In the non-dominant eye projection, clustered mAZ inputs were ∼55% closer together compared to isolated mAZ inputs (Fig. 5B, left panel). Once segregation was complete at P8, distances between clustered and isolated mAZ inputs were more similar (Fig. S5). In β2KO mice, distances between isolated and clustered mAZ inputs and their nearest clustered mAZ neighbor did not differ for either eye at P4 or P8 (Fig. 5/S5, right panels). These patterns are consistent with sAZ synapse addition when mAZ inputs are in close proximity during competition.

Clustered mAZ inputs are closer than isolated inputs during competition

Earlier version: (A-B) Distance between clustered and isolated mAZ inputs and the closest like-eye clustered mAZ input, shown for (A) CTB(+) and (B) CTB(-) projections at P4 in WT and β2KO mice. Boxes indicate the 25%-75% distribution of input measurements from N=3 biological replicates and whiskers extend to 1.5 times the interquartile range. Gray dots represent individual distance measurements for all mAZ inputs. Black and red dots represent mean values from separate biological replicates and black lines connect measurements within each replicate (N=3 for each age and genotype). Statistical significance was determined using a linear mixed model ANOVA with post-hoc Bonferroni correction. Black asterisks indicate significant differences. *p<0.05, ***p<0.001.

Revised version: (A-B) Distance between clustered and isolated mAZ inputs and the closest like-eye clustered mAZ input, shown for (A) CTB(+) and (B) CTB(-) projections at P4 in WT and β2KO mice. Boxes indicate the 25%-75% distribution of input measurements from N=3 biological replicates and whiskers extend to 1.5 times the interquartile range. Gray dots represent individual distance measurements for all mAZ inputs. Black and red dots represent mean values from separate biological replicates and black lines connect measurements within each replicate (N=3 for each age and genotype). Statistical significance was determined using a linear mixed model ANOVA with post-hoc Bonferroni correction. For each genotype, p-values were corrected for multiple testing with Benjamini-Hochberg FDR correction (α = 0.05) at each age. Black asterisks indicate significant differences. *p(adj)<0.05.

Finally, we tested whether sAZ vesicle pool volume varies with distance to mAZ inputs. We classified each sAZ as “near” (≤1.5 μm) or “far” (>1.5 μm) relative to the nearest like-eye mAZ input. Across ages and genotypes, vesicle pool volume was similar for near and far sAZ synapses in both eyes (Fig. S6). Thus, vesicle pool size in sAZ synapses appears independent of their proximity to mAZ inputs.

Discussion

Spontaneous retinal activity guides eye-specific refinement by strengthening dominant eye inputs in the correct territory and pruning weaker inputs from the competing eye. Using volumetric super-resolution microscopy and eye-specific synaptic immunolabeling, we previously identified activity-dependent, eye-specific differences in presynaptic vesicle pool size and vesicle association with the active zone (AZ), both favoring the dominant eye during synaptic competition 26. Here, we reanalyzed this dataset to measure synaptic spatial arrangements during normal development and how these are affected in a mutant model with abnormal eye-specific segregation.

Early input maturation and addition of release sites

Before eye-opening, retinal ganglion cell (RGC) axons from both eyes formed terminals with either multiple active zones (mAZ inputs) or a single active zone (sAZ synapses). Electron microscopy of dAPEX2-labeled ipsilateral RGC inputs revealed individual boutons with multiple active zones. Prior electron microscopy studies in mice showed no clustering from neighboring RGC boutons at postnatal day 8 10, 11, suggesting the mAZ inputs seen in STORM images are single boutons containing multiple release sites. During competition, the dominant eye projection generated more of these mAZ inputs, each with more active zones and larger vesicle pools than the non-dominant eye. We observed a linear relationship between presynaptic vesicle signal volume and active zone number, suggesting that release site addition and vesicle pool expansion may be linked during input maturation. Functional recordings indicate that release site addition is the primary driver of retinogeniculate input strengthening after eye-opening 24. Our findings reveal that eye-specific differences in release site addition at individual terminals emerge at the earliest stages of binocular refinement.

Spatial interactions during competition

Ipsilateral RGC axons are delayed in entering the dLGN until after contralateral inputs have already innervated the nucleus 55. Despite their late arrival, these axons competed by forming numerous sAZ inputs, which equalized the overall input density between the eyes in the future contralateral eye domain at postnatal day 4 (P4). Most of these synapses formed nearby other like-eye neighbors, and were enriched near mAZ inputs, supporting the idea of release site addition at strong synapses. Neighboring sAZ synapses could ultimately mature into mAZ inputs as release sites are added and vesicle pools expand.

Earlier version: Distance analysis showed that clustered mAZ boutons were closer together than isolated mAZ boutons, further supporting release site addition near strong neighboring inputs. It remains unknown, however, whether nearby mAZ inputs originate from the same RGC. Nearby mAZ inputs within individual RGC arbors could promote nearby release site addition via intrinsic mechanisms, such as increased presynaptic calcium entry or surface delivery of synaptogenic molecules. Boutons with multiple release sites could also trigger non-cell-autonomous signaling mechanisms that stabilize neighboring inputs from the same eye 48, 53, 54. Release site addition increases extracellular glutamate concentration and promotes spillover onto adjacent synapses, which enhances the excitability of developing relay cells 25 and could contribute to long term synaptic plasticity evoked by high frequency spiking in retinal wave bursts 56–59. Consequently, release site addition and local bouton clustering may act in tandem to stabilize coactive inputs.

Revised version: Distance analysis showed that clustered mAZ boutons were closer together than isolated mAZ boutons, further supporting release site addition near strong neighboring inputs. It remains unknown, however, whether nearby mAZ inputs originate from the same RGC. Nearby mAZ inputs within individual RGC arbors could promote release site addition via intrinsic mechanisms that scale with release site number, such as increased presynaptic calcium entry or surface delivery of synaptogenic molecules via presynaptic release. Boutons with multiple release sites could also trigger non-cell-autonomous signaling mechanisms that stabilize neighboring inputs from the same eye 48, 53, 54. Release site addition increases extracellular glutamate concentration and promotes spillover onto adjacent synapses, which enhances the excitability of developing relay cells 25 and could contribute to long term synaptic plasticity evoked by high frequency spiking in retinal wave bursts 56–59. Consequently, release site addition and local bouton clustering may act in tandem to stabilize coactive inputs.

Competitive refinement also involves synapse elimination and axonal retraction through punishment signals. Genetic deletions of VGluT2 or RIM1 proteins in ipsilaterally-projecting RGCs decreased presynaptic vesicle release and prevented retraction of contralateral RGC axons from the ipsilateral territory 39, 45. One downstream mediator of synaptic punishment is JAK2 kinase, which is phosphorylated in less active synapses 52. Similar to neurotransmission mutant phenotypes, disruption of JAK2 signaling prevents axon retraction (punishment) during competition 52. The spatial analyses we developed here will enable future mapping of input-specific punishment signals during synaptic competition. This includes phospho-JAK2 as well as molecular tags for glial pruning of weak inputs during eye-specific segregation 60–62.

Requirement for spontaneous retinal activity

Earlier version: β2KO mice showed significant defects in synapse addition during the first postnatal week. mAZs still formed with similar overall ratios as in wild type controls (Fig. S1, Table S3) and eye-specific differences in vesicle pool size still emerged. However, β2KO mice failed to form many sAZ synapses at the height of competition, particularly in the late-arriving non-dominant (ipsilateral) eye projection. The failure to add synapses reduced synaptic clustering and more inputs formed in isolation in the mutants compared to controls.

Revised version: β2KO mice showed significant defects in synapse addition during the first postnatal week. mAZs still formed with similar overall ratios as in wild type controls (Fig. S1, Table S3) and eye-specific differences in vesicle pool size still emerged. However, β2KO mice failed to form many sAZ synapses at the height of competition (P4), particularly in the late-arriving non-dominant (ipsilateral) eye projection. The failure to add synapses could explain the observation that synaptic clustering was reduced and more inputs formed in isolation in the mutants compared to controls.

Earlier version: While our findings support a role for spontaneous retinal activity in presynaptic release site addition and clustering, activity-dependent postsynaptic mechanisms also influence input refinement at later stages. Retinogeniculate synapses undergo postsynaptic strengthening and weakening through potentiation and depression mediated by AMPARs and NMDARs 56–59. After eye-specific segregation, spontaneous retinal activity is required for postsynaptic AMPAR insertion, synaptic strengthening, and elimination of weaker inputs 63. Continued maintenance of segregation depends on calcium influx into relay neurons via L-type calcium channels, further implicating postsynaptic signaling in late-stage refinement 64, 65. Input targeting may also be guided by molecular cues to form non-random, eye-specific connections with postsynaptic targets. Ipsilaterally-projecting RGCs have distinct gene expression profiles that specify axon guidance and may further support eye-specific synaptic targeting 66, 67. Spontaneous retinal activity may permit axons to read out molecular regulators of synaptogenesis, as previously shown for RGC axon retraction 68.

Revised version: While our results highlight developmental changes in presynaptic release site addition and clustering, activity-dependent postsynaptic mechanisms also influence input refinement at later stages. Retinogeniculate synapses undergo postsynaptic strengthening and weakening through potentiation and depression mediated by AMPARs and NMDARs 56–59. After eye-specific segregation, spontaneous retinal activity is required for postsynaptic AMPAR insertion, synaptic strengthening, and elimination of weaker inputs 63. Continued maintenance of segregation depends on calcium influx into relay neurons via L-type calcium channels, further implicating postsynaptic signaling in late-stage refinement 64, 65. Input targeting may also be guided by molecular cues to form non-random, eye-specific connections with postsynaptic targets. Ipsilaterally-projecting RGCs have distinct gene expression profiles that specify axon guidance and may further support eye-specific synaptic targeting 66, 67. Spontaneous retinal activity may permit axons to read out molecular regulators of synaptogenesis, as previously shown for RGC axon retraction 68.

Release site addition as a general mechanism underlying synaptic competition

Earlier version: Early synaptic clustering during retinogeniculate development resembles other circuits where neural activity guides competitive synaptic and axonal remodeling. At the neuromuscular junction (NMJ), motor neuron terminals compete for control of a postsynaptic muscle fiber; a single motor axon input strengthens while competing axons are eliminated 69–71. Competition at the NMJ depends on inter-synaptic distance, with motor axons losing connections located near stronger competing synapses 69, 70. As winning terminals are enlarged, presynaptic release site addition maintains a consistent density of Bassoon clusters (∼2-3/μm2) 72. Neighboring inputs with weaker synaptic transmission are selectively eliminated 73, 74. Competing motor axons differ in their release probabilities early in development 75, suggesting that presynaptic efficacy triggers local “stabilization” signals in winners and “punishment” signals in losers 76.

Revised version: Early synaptic clustering during retinogeniculate development resembles other circuits where neural activity guides competitive synaptic and axonal remodeling. At the neuromuscular junction (NMJ), motor neuron terminals compete for control of a postsynaptic muscle fiber; a single motor axon input strengthens while competing axons are eliminated 69–71. Competition at the NMJ depends on inter-synaptic distance, with motor axons losing connections located near stronger competing synapses 69, 70. As winning terminals are enlarged, presynaptic release site addition maintains a consistent density of Bassoon clusters (∼2-3/μm2) 72. Neighboring inputs with weaker synaptic transmission are selectively eliminated 73, 74. Competing motor axons differ in their release probabilities early in development 75, suggesting that presynaptic efficacy triggers local “stabilization” signals in winners and “punishment” signals in losers 76. Presynaptic agrin release stabilizes postsynaptic receptor clusters by counteracting the destabilizing, dispersal effects of acetylcholine release, emphasizing the importance of early presynaptic transmission in postsynaptic stabilization 77, 78.

Removed in the revised version: Presynaptic agrin release stabilizes postsynaptic receptor clusters by counteracting the destabilizing, dispersal effects of acetylcholine release, emphasizing the importance of early presynaptic transmission in postsynaptic stabilization 77, 78.

Removed in the revised version: Similarly, in the developing cerebellum, Purkinje cells initially receive inputs from multiple climbing fibers (CFs) during the first postnatal week. Subsequently, one winning CF emerges and consolidates its synaptic inputs, while losing CFs are pruned 79, 80.

Earlier version: Here again, the winning input strengthens by adding presynaptic release sites, which increases multivesicular release and elevates glutamate concentration in the synaptic cleft 81–83. Postsynaptic ultrastructural and molecular changes occur several days later as dominant CF inputs potentiate and losing CF inputs depress 82, 84, 85 through spike-timing dependent remodeling 86, 87. Across these models, the earliest distinguishing feature between competing inputs is relative presynaptic transmission strength. Thus, release-site addition may be a conserved mechanism that biases synaptic refinement outcomes across developing neural circuits.

Revised version: Similarly, in the developing cerebellum, Purkinje cells initially receive inputs from multiple climbing fibers (CFs) during the first postnatal week. Subsequently, one winning CF emerges and consolidates its synaptic inputs, while losing CFs are pruned 79, 80. Here again, the winning input strengthens by adding presynaptic release sites, which increases multivesicular release and elevates glutamate concentration in the synaptic cleft 81–83. Postsynaptic ultrastructural and molecular changes occur several days later as dominant CF inputs potentiate and losing CF inputs depress 82, 84, 85 through spike-timing dependent remodeling 86, 87. Across these models, the earliest distinguishing feature between competing inputs is relative presynaptic transmission strength. Thus, release-site addition may be a conserved mechanism that biases synaptic refinement outcomes across developing neural circuits.

Materials and methods

The raw imaging data in this paper were previously reported 26. Materials and methods below are adapted from this work. All Matlab and Python code used in the work is available on GitHub (https://github.com/SpeerLab). Raw STORM images of the full data are available on the open-access Brain Imaging Library 88. These images can be accessed here: https://doi.brainimagelibrary.org/doi/10.35077/g.1164.

Animals

Earlier version: Wild-type C57BL/6J mice (Stock Number 000664) and APEX2 mice (Stock Number 032764) used in this study were purchased from the Jackson Laboratory. ET33Cre mice and β2KO were generously gifted by Drs. Eric M. Ullian (University of California, San Francisco) and Michael C. Crair (Yale School of Medicine), respectively. All experimental procedures were performed in accordance with an animal study protocol approved by the Institutional Animal Care and Use Committee (IACUC) at the University of Maryland. Neonatal male and female mice were used interchangeably for all experiments. Tissue from biological replicates (N=3 animals) was collected for each age (P2/P4/P8) from each genotype (WT and β2KO) (18 animals total). Primers used for genotyping β2KO mice are: forward: CAGGCGTTATCCACAAAGACAGA; reverse:

Revised version: Wild-type C57BL/6J mice (Stock Number 000664) and APEX2 mice (Stock Number 032764) used in this study were purchased from the Jackson Laboratory. ET33Cre mice and β2KO were generously gifted by Drs. Eric M. Ullian (University of California, San Francisco) and Michael C. Crair (Yale School of Medicine), respectively. All experimental procedures were performed in accordance with an animal study protocol approved by the Institutional Animal Care and Use Committee (IACUC) at the University of Maryland. Neonatal male and female mice were used interchangeably for all experiments. Tissue from biological replicates (N=3 animals) was collected for each age (P2/P4/P8) from each genotype (WT and β2KO) (18 animals total). Primers used for genotyping β2KO mice are: forward: CAGGCGTTATCCACAAAGACAGA; reverse: TTGAGGGGAGCAGAACAGAATC; mutant reverse: ACTTGGGTTTGGGCGTGTTGAG. Primers used for genotyping ET33Cre mice are: Forward: GGTCCTTGGCAGATGGGCAT; reverse: CGGCAAACGGACAGAAGCATT.

Removed in the revised version: TTGAGGGGAGCAGAACAGAATC; mutant reverse: ACTTGGGTTTGGGCGTGTTGAG.

Removed in the revised version: Primers used for genotyping ET33Cre mice are: Forward: GGTCCTTGGCAGATGGGCAT; reverse: CGGCAAACGGACAGAAGCATT.

Electron microscopy tissue preparation

Animals were deeply anesthetized with ketamine/xylazine and transcardially perfused with 5-10 mls of 37°C 0.9% sterile saline (pH 7.2) followed by 20-30 mls of 37°C 4% EM-Grade paraformaldehyde (PFA, Electron Microscopy Sciences), 2% EM-Grade glutaraldehyde (GA, Electron Microscopy Sciences), 4mM calcium chloride (CaCl2, Sigma-Aldrich) in 0.2M cacodylate buffer (pH 7.4). Brains were postfixed in the same perfusion fixative solution overnight at 4°C. Brains were vibratome sectioned in 0.2M cacodylate buffer (pH 7.4) at 100 μm. A circular tissue punch (∼500 μm diameter) containing the dLGN was microdissected from each section using a blunt-end needle. dLGN sections were washed in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator (4x 20 minutes each). Sections were incubated in 20mM glycine (Sigma-Aldrich) in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator for 30 minutes, followed by 5 x 20 minutes washes in 0.2M cacodylate buffer (pH 7.4) at 4°C on a rotator. dLGN sections were immersed in 0.5mg/mL DAB solution, covered with foil, for 30 minutes at 4°C on a rotator. 10 μL of 0.03% hydrogen peroxide (Sigma-Aldrich) was mixed into the DAB solution (1 mL) and incubated for 10 minutes at 4°C on a rotator (light protected). The reaction was quenched with 0.2M cacodylate buffer (pH 7.4) washes (3 x 1-minute washes followed by 2 x 20 minutes washes).

Tissue preparation for scanning electron microscopy

dLGN sections were fixed in 2% osmium tetroxide (Electron Microscopy Sciences) in 0.15M cacodylate buffer (pH 7.4) for 45 minutes at room temperature. Sections were reduced with 2.5% potassium ferricyanide (Electron Microscopy Sciences) in 0.15M cacodylate buffer (pH 7.4) for 45 minutes at room temperature in the dark, followed by 2 x 10-minute washes in double distilled water. Sections were incubated in 1% aqueous thiocarbohydrazide (Electron Microscopy Sciences) for 20 minutes, followed by 2 x 10-minute washes in double distilled water. Samples were fixed in 1% aqueous osmium tetroxide for 45 minutes at room temperature, followed by 2 x 10-minute washes in double distilled water. Sections were postfixed in 1% uranyl acetate (Electron Microscopy Sciences) in 25% ethanol in the dark for 20 minutes at room temperature and washed and stored in double distilled water overnight. The next day, samples were stained with aqueous lead aspartate at 60°C for 30 minutes and washed with double distilled water 2 x 10 minutes. Tissues were dehydrated in a graded dilution series of 100% ethanol (35%; 50%; 70%; 95%; 100%; 100%; 100% EtOH) for 10 minutes each at room temperature. Samples were immersed in propylene oxide (Electron Microscopy Sciences) 3 x 10 minutes at room temperature, followed by a series of epon resin/propylene oxide (812 Epon Resin, Electron Microscopy Sciences) exchanges with increasing resin concentrations (50% resin/50% propylene oxide; 65% resin/35% propylene oxide; 75% resin/25% propylene oxide; 100% resin; 100% resin) for 90 minutes each. Tissues were transferred to BEEM capsules (Electron Microscopy Sciences) that were filled with 100% resin and polymerized for at least 48 hours at 60°C.

Transmission electron microscopy image acquisition

Plasticized sections were cut at 70 nm with a Histo Jumbo diamond knife (DiATOME) using a Leica UC7 ultramicrotome. Sections were decompressed using chloroform vapor and collected onto 3.05mm 200 mesh Gilder thin bar hexagonal mesh copper grids (T200H-Cu, Electron Microscopy Sciences). Sections were imaged unstained on a Hitachi HT7700 transmission electron microscope (HT7700, Hitachi High-Tech America, Inc.) at 80kV.

Eye injections

Intraocular eye injections were performed one day before tissue collection. Briefly, mice were anesthetized by inhalant isoflurane and sterile surgical spring scissors were used to gently part the eyelid to expose the corneoscleral junction. A small hole was made in the eye using a sterile 34-gauge needle and ∼0.5 μL of cholera toxin subunit B conjugated with Alexa Fluor 488 (CTB-488, ThermoFisher Scientific, Catalogue Number: C34775) diluted in 0.9% sterile saline was intravitreally pressure-injected into the right eye using a pulled-glass micropipette coupled to a Picospritzer (Parker Hannifin).

dLGN tissue preparation for STORM imaging

Removed in the revised version: Animals were deeply anesthetized with ketamine/xylazine and transcardially perfused with 5-10 mls of 37°C 0.9% sterile saline followed by 10 mls of room temperature 4% EM-Grade paraformaldehyde (PFA, Electron Microscopy Sciences) in 0.9% saline.

Earlier version: Brains were embedded in 2.5% agarose and vibratome sectioned in the coronal plane at 100 μm. From the full anterior-posterior series of dLGN sections (∼6-8 sections) we selected the central two sections for staining in all biological replicates. These sections were morphologically consistent with Figs. 134-136 (5.07-5.31 mm) of the postnatal day 6 mouse brain from Paxinos’s “Atlas of the developing mouse brain” Academic Press, 2020 89. Selected sections were postfixed in 4% PFA for 30 minutes at room temperature and washed for 30-40 minutes in 1X PBS. The dLGN was identified by the presence of CTB-488 signals using a fluorescence dissecting microscope. A circular tissue punch (∼500 μm diameter) containing the dLGN was microdissected from each section using a blunt-end needle. A small microknife cut was made at the dorsal edge of the dLGN which, together with the CTB-488 signal, enabled us to identify the dLGN orientation during image acquisition.

Revised version: Animals were deeply anesthetized with ketamine/xylazine and transcardially perfused with 5-10 mls of 37°C 0.9% sterile saline followed by 10 mls of room temperature 4% EM-Grade paraformaldehyde (PFA, Electron Microscopy Sciences) in 0.9% saline. Brains were embedded in 2.5% agarose and vibratome sectioned in the coronal plane at 100 μm. From the full anterior-posterior series of dLGN sections (∼6-8 sections) we selected the central two sections for staining in all biological replicates. These sections were morphologically consistent with Figs. 134-136 (5.07-5.31 mm) of the postnatal day 6 mouse brain from Paxinos’s “Atlas of the developing mouse brain” Academic Press, 2020 89. Selected sections were postfixed in 4% PFA for 30 minutes at room temperature and washed for 30-40 minutes in 1X PBS. The dLGN was identified by the presence of CTB-488 signals using a fluorescence dissecting microscope. A circular tissue punch (∼500 μm diameter) containing the dLGN was microdissected from each section using a blunt-end needle. A small microknife cut was made at the dorsal edge of the dLGN which, together with the CTB-488 signal, enabled us to identify the dLGN orientation during image acquisition.

Immunohistochemistry

We used a serial-section single-molecule localization imaging approach to prepare samples and collect super-resolution fluorescence imaging volumes as previously described 26. dLGN tissue punches were blocked in 10% normal donkey serum (Jackson ImmunoResearch, Catalogue Number: 017-000-121) with 0.3% Triton X-100 (Sigma-Aldrich Inc.) and 0.02% sodium azide (Sigma-Aldrich Inc.) diluted in 1X PBS for 2-3 hours at room temperature and then incubated in primary antibodies for ∼72 hours at 4°C. Primary antibodies used were Rabbit anti-Homer1 (Synaptic Systems, Catalogue Number: 160003, 1:100) to label postsynaptic densities (PSDs), mouse anti-Bassoon (Abcam, Catalogue Number AB82958, 1:100) to label presynaptic active zones (AZs), and guinea pig anti-VGluT2 (Millipore, Catalogue Number AB251-I, 1:100) to label presynaptic vesicles. Following primary antibody incubation, tissues were washed in 1X PBS for 6 x 20 minutes at room temperature and incubated in secondary antibody solution overnight for ∼36 hours at 4°C. The secondary antibodies used were donkey anti-rabbit IgG (Jackson ImmunoResearch, Catalogue Number 711-005-152, 1:100) conjugated with Dy749P1 (Dyomics, Catalogue Number 749P1-01) and Alexa Fluor 405 (ThermoFisher, Catalogue Number: A30000), donkey anti-mouse IgG (Jackson ImmunoResearch, Catalogue Number 715-005-150, 1:100) conjugated with Alexa Fluor 647 (ThermoFisher, Catalogue Number: A20006) and Alexa Fluor 405, and donkey anti-guinea pig IgG (Jackson ImmunoResearch, Catalogue Number 706-005-148, 1:100) conjugated with Cy3b (Cytiva, Catalogue Number: PA63101). Tissues were washed 6 x 20 minutes in 1X PBS at room temperature after secondary antibody incubation.

Postfixation, dehydration, and embedding in epoxy resin

Tissue embedding was performed as previously described 26. Tissues were postfixed with 3% PFA + 0.1% GA (Electron Microscopy Sciences) in PBS for 2 hours at room temperature and then washed in 1X PBS for 20 minutes. To plasticize the tissues for ultrasectioning, the tissues were first dehydrated in a graded dilution series of 100% ethanol (50%/70%/90%/100%/100% EtOH) for 15 minutes each at room temperature and then immersed in a series of epoxy resin/100% EtOH exchanges (Electron Microscopy Sciences) with increasing resin concentration (25% resin/75% ethanol; 50% resin/50% ethanol; 75% resin/25% ethanol; 100% resin; 100% resin) for 2 hours each. Tissues were transferred to BEEM capsules (Electron Microscopy Sciences) that were filled with 100% resin and polymerized for 16 hours at 70°C.

Ultrasectioning

Plasticized tissue sections were cut using a Leica UC7 ultramicrotome at 70 nm using a Histo Jumbo diamond knife (DiATOME). Chloroform vapor was used to reduce compression after cutting. For each sample, ∼100 sections were collected on a coverslip coated with 0.5% gelatin and 0.05% chromium potassium (Sigma-Aldrich Inc.), dried at 60 degrees for 25 minutes, and protected from light prior to imaging.

Imaging chamber preparation

Coverslips were chemically etched in 10% sodium ethoxide for 5 minutes at room temperature to remove the epoxy resin and expose the dyes to the imaging buffer for optimal photoswitching. Coverslips were then rinsed with ethanol and dH2O. To create fiducial beads for flat-field and chromatic corrections, we mixed 715/755nm and 540/560nm, carboxylate-modified microspheres (Invitrogen, Catalogue Numbers F8799 and F8809, 1:8 ratio respectively) to create a high-density fiducial marker and then further diluted the mixture at 1:750 with Dulbecco’s PBS to create a low-density bead solution. Both high- and low-density bead solutions were spotted on the coverslip (∼0.7 µL each) for flat-field and chromatic aberration correction respectively. Excess beads were rinsed away with dH2O for 1-2 minutes. The coverslip was attached to a glass slide with double-sided tape to form an imaging chamber. The chamber was filled with STORM imaging buffer (10% glucose, 17.5µM glucose oxidase, 708nM catalase, 10mM MEA, 10mM NaCl, and 200mM Tris) and sealed with epoxy.

Imaging setup

Imaging was performed using a custom single-molecule super-resolution imaging system. The microscope contained low (4x/10x air) and high (60x 1.4NA oil immersion) magnitude objectives mounted on a commercial frame (Nikon Ti-U) with back optics arranged for oblique incident angle illumination. We used continuous-wave lasers at 488nm (Coherent), 561nm (MPB), 647nm (MPB), and 750nm (MPB) to excite Alexa 488, Cy3B, Alexa 647, and Dy749P1 dyes respectively. A 405 nm cube laser (Coherent) was used to reactivate Dy749P1 and Alexa647 dye photoswitching. The microscope was fitted with a custom pentaband/pentanotch dichroic filter set and a motorized emission filter wheel. The microscope also contained an IR laser-based focus lock system to maintain optimal focus during automatic image acquisition. Images were collected on 640*640-pixel region of an sCMOS camera (ORCA-Flash4.0 V3, Hamamatsu Photonics) with a pixel size of ∼155 nm.

Automated image acquisition

Fiducials and tissue sections on the coverslip were imaged using the low magnification objective (4X) to create a mosaic overview of the specimen. Beads/sections were then imaged at high magnification (60X) to select regions of interest (ROIs) in the Cy3B and Alexa 488 channels. Before final image acquisition, laser intensities and the incident angle were adjusted to optimize photoswitching for STORM imaging and utilize the full dynamic range of the camera for conventional imaging.

Low-density bead images were taken in 16 partially overlapping ROIs. 715/755nm beads were excited using 750 nm light and images were collected through Dy749P1 and Alexa 647 emission filters. 540/560nm beads were excited using a 488 nm laser and images were collected through Alexa 647, Cy3B, and Alexa 488 emission filters. These fiducial images were later used to generate a non-linear warping transform to correct chromatic aberration. Next, ROIs within each tissue section were imaged at conventional (diffraction-limited) resolution in all four-color channels sequentially.

Following conventional image acquisition, a partially overlapping series of 9 images were collected in the high-density bead field for all 4 channels (Dy749P1, Alexa 647, Cy3B, and Alexa 488). These images were later used to perform a flat-field image correction of non-uniform laser illumination across the ROIs. Another round of bead images was taken as described above in a different ROI of the low-density bead field. These images were later used to confirm the stability of chromatic offsets during imaging. All ROIs within physical sections were then imaged by STORM for Dy749P1 and Alexa 647 channels. Images were acquired using a custom progression of increasing 405nm laser intensity to control single-molecule switching. 8000 frames of Dy749P1 channel images were collected (60 Hz imaging) followed by 12000 frames of Alexa 647 channel images (100 Hz). In a second imaging pass, the same ROIs were imaged for Cy3B and Alexa 488 channels, each for 8000 frames (60 Hz).

We imaged the ipsilateral and contralateral ROIs separately in each physical section of the dLGN. For consistency of ROI selection across biological replicates at each age, we identified the dorsal-ventral (DV) axis of the dLGN and selected ROIs within the center (core region) at 2/5 (ipsilateral ROI) and 4/5 (contralateral ROI) of the full DV length

Image processing

Earlier version: Single-molecule localization was performed using a previously described DAOSTORM algorithm modified for use with sCMOS cameras 90, 91. Molecule lists were rendered as 8-bit images with 15.5 nm pixel size where each molecule is plotted as an intensity distribution with an area reflecting its localization precision. Low-density fiducial images were used for chromatic aberration correction. We localized 715/755 beads in Dy749P1 and Alexa 647 channels, and 540/560 beads in Alexa 647, Cy3B, and Alexa 488 channels. A third-order polynomial transform map was generated by matching the positions of each bead in all channels to the Alexa 647 channel. The average residual error of bead matching was <15 nm for all channels. The transform maps were applied to both 4-color conventional and STORM images. Conventional images were upscaled (by 10X) to match the STORM image size. The method to align serial sections was previously described 26. STORM images were first aligned to their corresponding conventional images by image correlation. To generate an aligned 3D image stack from serial sections, we normalized the intensity of all Alexa 488 images and used these normalized images to generate both rigid and elastic transformation matrices for all four-color channels of both STORM and conventional data. The final image stack was then rotated and cropped to exclude incompletely imaged edge areas. To further confirm that the processed region corresponds to the contralateral dLGN, conventional CTB(+) signals of the labeled contralateral projection were thresholded to create a polygonal mask (a convex hull analysis linking the outermost CTB signals in the image volume).

Revised version: Single-molecule localization was performed using a previously described DAOSTORM algorithm modified for use with sCMOS cameras 90, 91. Molecule lists were rendered as 8-bit images with 15.5 nm pixel size where each molecule is plotted as an intensity distribution with an area reflecting its localization precision. Low-density fiducial images were used for chromatic aberration correction. We localized 715/755 beads in Dy749P1 and Alexa 647 channels, and 540/560 beads in Alexa 647, Cy3B, and Alexa 488 channels. A third-order polynomial transform map was generated by matching the positions of each bead in all channels to the Alexa 647 channel. The average residual error of bead matching was <15 nm for all channels. The transform maps were applied to both 4-color conventional and STORM images. Conventional images were upscaled (by 10X) to match the STORM image size. The method to align serial sections was previously described 26. STORM images were first aligned to their corresponding conventional images by image correlation. To generate an aligned 3D image stack from serial sections, we normalized the intensity of all Alexa 488 images and used these normalized images to generate both rigid and elastic transformation matrices for all four-color channels of both STORM and conventional data. The final image stack was then rotated and cropped to exclude incompletely imaged edge areas. To further confirm that the processed region corresponds to the contralateral dLGN, conventional CTB(+) signals of the labeled contralateral projection were thresholded to create a polygonal mask (a convex hull analysis linking the outermost CTB signals in the image volume). The mask was then applied to STORM images to exclude peripheral areas where CTB signals were absent or faint.

Removed in the revised version: The mask was then applied to STORM images to exclude peripheral areas where CTB signals were absent or faint.

Cell body filter

Removed in the revised version: The aligned STORM images had non-specific labeling of cell bodies in Dy749P1 and Alexa 647 channels corresponding to Homer1 and Bassoon immunolabels. To limit synaptic cluster identification to the neuropil region we identified cell bodies based on their Dy749P1 signal and excluded these regions from further image processing.

Earlier version: STORM images were convolved with a Gaussian function (σ=140 nm) and then binarized using the lower threshold of a two-level Otsu threshold method. We located connected components in the thresholded images and generated a mask based on components larger than e11 voxels. Because cell body clusters were orders of magnitude larger than synaptic clusters, the cell body filter algorithm was robust to a range of size thresholds. The mask was applied to images of all channels to exclude cell body areas.

Revised version: The aligned STORM images had non-specific labeling of cell bodies in Dy749P1 and Alexa 647 channels corresponding to Homer1 and Bassoon immunolabels. To limit synaptic cluster identification to the neuropil region we identified cell bodies based on their Dy749P1 signal and excluded these regions from further image processing. STORM images were convolved with a Gaussian function (σ=140 nm) and then binarized using the lower threshold of a two-level Otsu threshold method. We located connected components in the thresholded images and generated a mask based on components larger than e11 voxels. Because cell body clusters were orders of magnitude larger than synaptic clusters, the cell body filter algorithm was robust to a range of size thresholds. The mask was applied to images of all channels to exclude cell body areas.

Eye-specific synapse identification and quantification

Removed in the revised version: To correct for minor variance in image intensity across physical sections, we normalized the pixel intensity histogram of each section to the average histogram of all sections.

Earlier version: Image histograms were rescaled to make full use of the 8-bit range. Using a two-level Otsu threshold method, the conventional images were thresholded into three classes: a low-intensity background, low-intensity signals above the background representing non-synaptic labeling, and high-intensity signals representing synaptic structures. The conventional images were binarized by the lower two-level Otsu threshold, generating a mask for STORM images to filter out background signals. STORM images were convolved with a Gaussian function (σ= 77.5 nm) and thresholded using the higher two-level Otsu threshold. Following thresholding, connected components were identified in three dimensions using MATLAB ‘conncomp’ function. A watershedding approach was applied to split large clusters that were improperly connected. Clusters were kept for further analysis only if they contained aligned image information across two or more physical sections. We also removed all edge synapses from our analysis by excluding synapses that did not have blank image data on all adjacent sides.

Revised version: To correct for minor variance in image intensity across physical sections, we normalized the pixel intensity histogram of each section to the average histogram of all sections. Image histograms were rescaled to make full use of the 8-bit range. Using a two-level Otsu threshold method, the conventional images were thresholded into three classes: a low-intensity background, low-intensity signals above the background representing non-synaptic labeling, and high-intensity signals representing synaptic structures. The conventional images were binarized by the lower two-level Otsu threshold, generating a mask for STORM images to filter out background signals. STORM images were convolved with a Gaussian function (σ= 77.5 nm) and thresholded using the higher two-level Otsu threshold. Following thresholding, connected components were identified in three dimensions using MATLAB ‘conncomp’ function. A watershedding approach was applied to split large clusters that were improperly connected. Clusters were kept for further analysis only if they contained aligned image information across two or more physical sections. We also removed all edge synapses from our analysis by excluding synapses that did not have blank image data on all adjacent sides.

Removed in the revised version: To distinguish non-specific immunolabeling from true synaptic signals, we quantified two parameters for each cluster: cluster volume and cluster signal density calculated by the ratio of within-cluster pixels with positive signal intensity in the raw STORM images.

Earlier version: Two separate populations were identified in 2D histograms plotted from these two parameters. We manually selected the population with higher volumes and signal densities representing synaptic structures. To test the robustness of the manual selection, we performed multiple repeated measurements of the same data and discovered a between-measurement variance of <1% (data not shown).

Revised version: To distinguish non-specific immunolabeling from true synaptic signals, we quantified two parameters for each cluster: cluster volume and cluster signal density calculated by the ratio of within-cluster pixels with positive signal intensity in the raw STORM images. Two separate populations were identified in 2D histograms plotted from these two parameters. We manually selected the population with higher volumes and signal densities representing synaptic structures. To test the robustness of the manual selection, we performed multiple repeated measurements of the same data and discovered a between-measurement variance of <1% (data not shown).

To identify paired pre- and postsynaptic clusters, we first measured the centroid-centroid distance of each cluster in the Dy749P1 (Homer1) and Alexa 647 (Bassoon) channels to the closest cluster in the other channel. We next quantified the signal intensity of each opposing synaptic channel within a 140 nm shell surrounding each cluster. A 2D histogram was plotted based on the measured centroid-centroid distances and opposing channel signal densities of each cluster. Paired clusters with closely positioned centroids and high intensities of apposed channel signal were identified using the OPTICS algorithm. Retinogeniculate synapses were identified by pairing Bassoon (Alexa 647) clusters with VGluT2 (Cy3B) clusters using the same method as pre/post-synaptic pairing. Synapses from the right eye were identified by pairing VGluT2 clusters with CTB (Alexa 488) clusters. The volume of each cluster reflected the total voxel volume of all connected voxels, and the total signal intensity was a sum of voxel intensity within the volume of the connected voxels.

Multi-AZ synapse identification and quantification

Earlier version: To determine whether an eye-specific VGluT2 cluster is a mAZ synapse or a sAZ synapse, we measured the number of active zones (defined by individual Bassoon clusters) associated with each VGluT2 cluster in the dataset. A 3D shell was extended 140 nm from the surface voxels of each VGluT2 cluster and any Bassoon clusters that fell within the shell were considered to be associated with the target VGluT2 cluster.

Revised version: To determine whether an eye-specific VGluT2 cluster is a mAZ synapse or a sAZ synapse, we measured the number of active zones (defined by individual Bassoon clusters) associated with each VGluT2 cluster in the dataset. A 3D shell was extended 140 nm from the surface voxels of each VGluT2 cluster and any Bassoon clusters that fell within the shell were considered to be associated with the target VGluT2 cluster. The number of active zones (AZs) associated with each VGluT2 cluster was then measured. VGluT2 clusters associated with more than 1 AZ were defined as mAZ synapses, while those associated with only 1 AZ were defined as sAZ synapses.

Removed in the revised version: The number of active zones (AZs) associated with each VGluT2 cluster was then measured. VGluT2 clusters associated with more than 1 AZ were defined as mAZ synapses, while those associated with only 1 AZ were defined as sAZ synapses.

Quantification of mAZ and sAZ synapse VGluT2 cluster volume was performed using the “regionprops” function in MATLAB, which provided the voxel size and weighted centroid of each VGluT2 cluster. The search for sAZ synapses adjacent to mAZ synapses (synaptic clustering analysis) was conducted using a similar search approach as for associated Bassoon clusters, with expansion shell sizes ranging from 1 μm to 4 μm from the surface voxels of each mAZ synapse. The main Figs. in the study utilized an expansion distance of 1.5 μm. An eye-specific sAZ synapse was considered to be near a mAZ synapse if its weighted centroid fell within the expanded volume.

Quantification and statistical analysis

Removed in the revised version: Statistical analysis was performed using SPSS. Plots were generated by SPSS or R (ggplot2). Statistical details are presented in the figure legends and supplemental tables S2/S3. For all measurements in this paper, we analyzed N = 3 biological replicates (individual mice) for each genotype (WT and β2KO) at each age (P2, P4, and P8).

Earlier version: Cluster densities, synapse AZ number, average VGluT2 cluster volume, and all fractional measurements were presented as mean ± SEM values in paired bar graphs and statistical analysis was performed using two-tailed paired-T tests. Nonparametric Kolmogorov-Smirnov tests were used in all cumulative histogram comparisons. We used a linear mixed model to compare VGluT2 cluster volumes (Figs. 3 and S2) and distance measurements (Figs. 5 and S5). For VGluT2 cluster volume comparisons, the age or eye-of-origin was the fixed main factor and biological replicate IDs were nested random factors. In distance measurement comparisons, the mAZ synapse AZ number was the fixed main factor and biological replicate IDs were nested random factors. In violin plots, each violin showed the distribution of grouped data from all biological replicates from the same condition. Each black dot represents the median value of each biological replicate, and the horizontal black line represents the group median. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Asterisks in all figures indicate statistical significance: *P<0.05, **P<0.01, ***P<0.001. For VGluT2 volume comparisons, we calculated 5/95% confidence intervals based on the linear mixed model. 5/95% confidence intervals in synapse densities, AZ numbers, and distances were calculated for paired or unpaired data comparisons using SPSS.

Revised version: Statistical analysis was performed using SPSS. Plots were generated by SPSS or R (ggplot2). Statistical details are presented in the figure legends and supplemental tables S2/S3. For all measurements in this paper, we analyzed N = 3 biological replicates (individual mice) for each genotype (WT and β2KO) at each age (P2, P4, and P8). Cluster densities, synapse AZ number, average VGluT2 cluster volume, and all fractional measurements were presented as mean ± SEM values in paired bar graphs and statistical analysis was performed using two-tailed paired T-tests. Nonparametric Kolmogorov-Smirnov tests were used in all cumulative histogram comparisons. We used a linear mixed model to compare VGluT2 cluster volumes (Figs. 3 and S2) and distance measurements (Figs. 5 and S5). For VGluT2 cluster volume comparisons, the age or eye-of-origin was the fixed main factor and biological replicate IDs were nested random factors. In distance measurement comparisons, the mAZ synapse AZ number was the fixed main factor and biological replicate IDs were nested random factors. Effect sizes were calculated using Cohen’s d for parametric tests and epsilon squared for nonparametric tests. To address multiple testing, we applied false discovery rate (FDR) correction using the Benjamini-Hochberg method with α = 0.05 separately within each experimental condition (P2-WT, P2-ß2KO, P4-WT, P4-ß2KO, P8-WT, P8-ß2KO). 20-34 measurements (varying by age and genotype) were corrected within each of the six experimental conditions, resulting in condition-specific FDR-adjusted p-values presented in Table S2. In violin plots, each violin showed the distribution of grouped data from all biological replicates from the same condition. Each black dot represents the median value of each biological replicate, and the horizontal black line represents the group median. Black lines connect measurements of CTB(+) and CTB(-) populations from the same biological replicate. Asterisks in all figures indicate statistical significance after FDR correction: *p(adj)<0.05. For VGluT2 volume comparisons, we calculated 5/95% confidence intervals based on the linear mixed model. 5/95% confidence intervals in synapse densities, AZ numbers, and distances were calculated for paired or unpaired data comparisons using SPSS. All statistical results are listed in Table S2. All conclusions we draw in the main text from the data have corresponding confidence intervals listed in both the figure legends and Table S3.

Removed in the revised version: All statistical results are listed in Table S2. All conclusions we draw in the main text from the data have corresponding confidence intervals listed in both the figure legends and Table S3.